rule_files:
  - ../../monitoring/base/prometheus/alert_rules/metallb.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="metallb-controller"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: MetalLBControllerDown
        exp_alerts:
          - exp_labels:
              severity: error
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: MetalLB controller has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="metallb-speaker",instance="10.0.0.1",node="10.0.0.1"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: MetalLBSpeakerDown
        exp_alerts:
          - exp_labels:
              job: metallb-speaker
              severity: error
              node: 10.0.0.1
              instance: 10.0.0.1
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: MetalLB speaker has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="metallb-speaker",instance="10.0.0.1",node="10.0.0.1"}'
        values: '0+0x10'
      - series: 'kube_node_spec_taint{key="node.kubernetes.io/unreachable",effect="NoSchedule",node="10.0.0.1"}'
        values: '1+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: MetalLBSpeakerDown
        exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'up{job="metallb-speaker",instance="10.0.0.1",node="10.0.0.1"}'
        values: '1+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: MetalLBSpeakerDown
        exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'metallb_allocator_addresses_in_use_total{pool="a"}'
        values: '76+0x10'
      - series: 'metallb_allocator_addresses_in_use_total{pool="b"}'
        values: '75+0x10'
      - series: 'metallb_allocator_addresses_total{pool="a"}'
        values: '100+0x10'
      - series: 'metallb_allocator_addresses_total{pool="b"}'
        values: '100+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: MetalLBAddressPoolHighUtilization
        exp_alerts:
          - exp_labels:
              severity: warning
              pool: a
            exp_annotations:
              description: Address Pool of MetalLB will be exhausted.
              runbook: Please re-consider the address allocation planning.
