rule_files:
  - ../../monitoring/base/prometheus/alert_rules/teleport.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="teleport"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: TeleportDown
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: Teleport has disappeared from Prometheus target discovery.
