rule_files:
  - ../../monitoring/base/victoriametrics/rules/converted/etcd-alertrule.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'etcd_mvcc_db_total_size_in_bytes{instance="10.0.0.1:2379", job="etcd"}'
        values: '10+0x10 81+0x10'
      - series: 'etcd_server_quota_backend_bytes{instance="10.0.0.1:2379", job="etcd"}'
        values: '100+0x20'
    alert_rule_test:
      - eval_time: 20m
        alertname: DatabaseSpaceExceeded
        exp_alerts:
          - exp_labels:
              job: etcd
              instance: 10.0.0.1:2379
              severity: error
            exp_annotations:
              summary: "10.0.0.1:2379, etcd of etcd DB space uses more than 80%"
              runbook: "Please consider manual compaction and defrag. https://github.com/etcd-io/etcd/blob/master/Documentation/op-guide/maintenance.md"
  - interval: 1m
    input_series:
      - series: 'etcd_mvcc_db_total_size_in_bytes{instance="10.0.0.1:2379", job="etcd"}'
        values: '10+0x10 91+0x10'
      - series: 'etcd_server_quota_backend_bytes{instance="10.0.0.1:2379", job="etcd"}'
        values: '100+0x20'
    alert_rule_test:
      - eval_time: 20m
        alertname: DatabaseSpaceExceeded
        exp_alerts:
          - exp_labels:
              job: etcd
              instance: 10.0.0.1:2379
              severity: error
            exp_annotations:
              summary: "10.0.0.1:2379, etcd of etcd DB space uses more than 80%"
              runbook: "Please consider manual compaction and defrag. https://github.com/etcd-io/etcd/blob/master/Documentation/op-guide/maintenance.md"
          - exp_labels:
              job: etcd
              instance: 10.0.0.1:2379
              severity: error
            exp_annotations:
              summary: "10.0.0.1:2379, etcd of etcd DB space uses more than 90%"
              runbook: "Please consider manual compaction and defrag. https://github.com/etcd-io/etcd/blob/master/Documentation/op-guide/maintenance.md"
  - interval: 30m
    input_series:
      - series: 'etcd_mvcc_db_total_size_in_use_in_bytes{instance="10.0.0.1:2379", job="etcd"}'
        values: '0+15000001x2'
    alert_rule_test:
      - eval_time: 60m
        alertname: LogicalDatabaseUsageIncreaseRapidly
        exp_alerts:
          - exp_labels:
              job: etcd
              instance: 10.0.0.1:2379
              severity: warning
            exp_annotations:
              summary: "10.0.0.1:2379, etcd of etcd DB space increases 30MB/h"
              runbook: "Please consider to find root causes, and solve the problems"
  - interval: 1m
    input_series:
      - series: 'up{job="bootserver-etcd",instance="10.0.0.1"}'
        values: '0+0x10'
      - series: 'up{job="bootserver-etcd",instance="10.0.0.2"}'
        values: '0+0x10'
      - series: 'up{job="bootserver-etcd",instance="10.0.0.3"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: BootserverEtcdMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: BootserverEtcd has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="bootserver-etcd",instance="10.0.0.1"}'
        values: '1+0x10'
      - series: 'up{job="bootserver-etcd",instance="10.0.0.2"}'
        values: '0+0x10'
      - series: 'up{job="bootserver-etcd",instance="10.0.0.3"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: BootserverEtcdMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: BootserverEtcd has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="bootserver-etcd",instance="10.0.0.1"}'
        values: '1+0x10'
      - series: 'up{job="bootserver-etcd",instance="10.0.0.2"}'
        values: '1+0x10'
      - series: 'up{job="bootserver-etcd",instance="10.0.0.3"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 20m
        alertname: BootserverEtcdMissing
        exp_alerts:
          - exp_labels:
              severity: warning
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: BootserverEtcd has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="bootserver-etcd",instance="10.0.0.1"}'
        values: '1+0x10'
      - series: 'up{job="bootserver-etcd",instance="10.0.0.2"}'
        values: '1+0x10'
      - series: 'up{job="bootserver-etcd",instance="10.0.0.3"}'
        values: '1+0x10'
    alert_rule_test:
      - eval_time: 20m
        alertname: BootserverEtcdMissing
        exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'up{job="cke-etcd",instance="10.0.0.1"}'
        values: '0+0x10'
      - series: 'up{job="cke-etcd",instance="10.0.0.2"}'
        values: '0+0x10'
      - series: 'up{job="cke-etcd",instance="10.0.0.3"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: CKEEtcdMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: CKEEtcd has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="cke-etcd",instance="10.0.0.1"}'
        values: '1+0x10'
      - series: 'up{job="cke-etcd",instance="10.0.0.2"}'
        values: '0+0x10'
      - series: 'up{job="cke-etcd",instance="10.0.0.3"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: CKEEtcdMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: CKEEtcd has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="cke-etcd",instance="10.0.0.1"}'
        values: '1+0x10'
      - series: 'up{job="cke-etcd",instance="10.0.0.2"}'
        values: '1+0x10'
      - series: 'up{job="cke-etcd",instance="10.0.0.3"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 20m
        alertname: CKEEtcdMissing
        exp_alerts:
          - exp_labels:
              severity: warning
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: CKEEtcd has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="cke-etcd",instance="10.0.0.1"}'
        values: '1+0x10'
      - series: 'up{job="cke-etcd",instance="10.0.0.2"}'
        values: '1+0x10'
      - series: 'up{job="cke-etcd",instance="10.0.0.3"}'
        values: '1+0x10'
    alert_rule_test:
      - eval_time: 20m
        alertname: CKEEtcdMissing
        exp_alerts: []

  - interval: 1m
    input_series: []
    alert_rule_test:
      - eval_time: 3h
        alertname: CKEEtcdBackupNotExecuted
        exp_alerts:
          - exp_labels:
              severity: error
            exp_annotations:
              summary: Backup job for CKEEtcd is not executed.
              runbook: TBD
  - interval: 1m
    input_series:
      - series: 'cke_etcd_backup_last_executed_time{instance="10.0.0.1"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_executed_time{instance="10.0.0.2"}'
        values: '0+1x10800'
    alert_rule_test:
      - eval_time: 3h
        alertname: CKEEtcdBackupNotExecuted
        exp_alerts:
          - exp_labels:
              severity: error
            exp_annotations:
              summary: Backup job for CKEEtcd is not executed.
              runbook: TBD
  - interval: 1m
    input_series:
      - series: 'cke_etcd_backup_last_executed_time{instance="10.0.0.1"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_executed_time{instance="10.0.0.2"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_executed_time{instance="10.0.0.3"}'
        values: '0+1x10800'
    alert_rule_test:
      - eval_time: 3h
        alertname: CKEEtcdBackupNotExecuted
        exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'cke_etcd_backup_last_executed_time{instance="10.0.0.1"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_executed_time{instance="10.0.0.2"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_executed_time{instance="10.0.0.3"}'
        values: '0+0x10800'
    alert_rule_test:
      - eval_time: 3h
        alertname: CKEEtcdBackupNotExecuted
        exp_alerts:
          - exp_labels:
              severity: error
              instance: 10.0.0.3
            exp_annotations:
              summary: Backup job for CKEEtcd is not executed.
              runbook: TBD
  - interval: 1m
    input_series: []
    alert_rule_test:
      - eval_time: 3h
        alertname: CKEEtcdBackupNotSucceeded
        exp_alerts:
          - exp_labels:
              severity: error
            exp_annotations:
              summary: Backup job for CKEEtcd is not succeeded.
              runbook: TBD
  - interval: 1m
    input_series:
      - series: 'cke_etcd_backup_last_succeeded_time{instance="10.0.0.1"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_succeeded_time{instance="10.0.0.2"}'
        values: '0+1x10800'
    alert_rule_test:
      - eval_time: 3h
        alertname: CKEEtcdBackupNotSucceeded
        exp_alerts:
          - exp_labels:
              severity: error
            exp_annotations:
              summary: Backup job for CKEEtcd is not succeeded.
              runbook: TBD
  - interval: 1m
    input_series:
      - series: 'cke_etcd_backup_last_succeeded_time{instance="10.0.0.1"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_succeeded_time{instance="10.0.0.2"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_succeeded_time{instance="10.0.0.3"}'
        values: '0+1x10800'
    alert_rule_test:
      - eval_time: 3h
        alertname: CKEEtcdBackupNotSucceeded
        exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'cke_etcd_backup_last_succeeded_time{instance="10.0.0.1"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_succeeded_time{instance="10.0.0.2"}'
        values: '0+1x10800'
      - series: 'cke_etcd_backup_last_succeeded_time{instance="10.0.0.3"}'
        values: '0+0x10800'
    alert_rule_test:
      - eval_time: 3h
        alertname: CKEEtcdBackupNotSucceeded
        exp_alerts:
          - exp_labels:
              severity: error
              instance: 10.0.0.3
            exp_annotations:
              summary: Backup job for CKEEtcd is not succeeded.
              runbook: TBD
