rule_files:
  - ../../monitoring/base/victoriametrics/rules/converted/ingress-watcher-alertrule.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'ingresswatcher_http_get_successful_total{code="200", instance="1.2.3.4", job="ingress-watcher", path="http://hoge/"}'
        values: '0+5x15'
      - series: 'ingresswatcher_http_get_total{instance="1.2.3.4", job="ingress-watcher", path="http://hoge/"}'
        values: '0+6x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: IngressDown
        exp_alerts:
          - exp_labels:
              instance: 1.2.3.4
              job: ingress-watcher
              path: http://hoge/
              severity: critical
            exp_annotations:
              runbook: TBD
              summary: ingress http://hoge/ is down (from 1.2.3.4).
  - interval: 1m
    input_series:
      - series: 'ingresswatcher_update_time{instance="1.2.3.4", job="ingress-watcher"}'
        values: '10+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: IngressWatcherDown
        exp_alerts:
          - exp_labels:
              instance: 1.2.3.4
              job: ingress-watcher
              severity: critical
            exp_annotations:
              runbook: TBD
              summary: ingress watcher 1.2.3.4 is down.
