rule_files:
  - ../../monitoring/base/victoriametrics/rules/converted/ingress-alertrule.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="contour", kubernetes_namespace="ingress-global"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: ContourGlobalDown
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: Contour has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="contour-envoy", kubernetes_namespace="ingress-global"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: IngressGlobalDown
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems
              summary: Ingress has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'contour_httpproxy_invalid{kubernetes_namespace="ingress-global",namespace="ours",vhost="xxx"}'
        values: '1+0x10'
      - series: 'kube_namespace_labels{namespace="ours",label_team="neco"}'
        values: '1+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: InvalidHTTPProxyExist
        exp_alerts:
          - exp_labels:
              severity: error
              kubernetes_namespace: ingress-global
              namespace: ours
              vhost: xxx
              label_team: neco
            exp_annotations:
              summary: Invalid HTTPProxy resource exists.
              runbook: Please check the status of HTTPProxy.
  - interval: 1m
    input_series:
      - series: 'contour_httpproxy_invalid{kubernetes_namespace="ingress-global",namespace="others",vhost="xxx"}'
        values: '1+0x10'
      - series: 'kube_namespace_labels{namespace="ours",label_team="neco"}'
        values: '1+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: InvalidHTTPProxyExist
        exp_alerts: []
