rule_files:
  - ../../monitoring/base/victoriametrics/rules/converted/network-policy-alertrule.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="calico-node",instance="10.0.0.1:9091",node="10.0.0.1"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: CalicoNodeDown
        exp_alerts:
          - exp_labels:
              job: calico-node
              severity: critical
              node: 10.0.0.1
              instance: 10.0.0.1:9091
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems.
              summary: CalicoNode has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="calico-node",instance="10.0.0.1:9091",node="10.0.0.1"}'
        values: '0+0x10'
      - series: 'kube_node_spec_taint{key="node.kubernetes.io/unreachable",effect="NoSchedule",node="10.0.0.1"}'
        values: '1+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: CalicoNodeDown
        exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'up{job="calico-node",instance="10.0.0.1:9091",node="10.0.0.1"}'
        values: '1+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: CalicoNodeDown
        exp_alerts: []
